# Un problema de KNN

En este taller estudiarás conceptos mostrados en el tutorial "Solución de problemas con KNN", utilizando el conjunto de datos correspondiente al desempeño y desgaste de empleados. Particularmente, realizarás los siguientes procesos:

1. Cargar un conjunto de datos.
2. Preparar los datos para el modelado.
3. Graficar una curva de validación para el hiperparámetro _k_.
4. Evaluar el mejor modelo resultante.

Entonces, dadas algunas características en áreas como la educación, trabajos previos, salario, entre otras, queremos clasificar a un empleado como agotado (con bajo desempeño) o no. Antes de iniciar, vamos a importar las librerías necesarias:

```python
import pandas as pd
import matplotlib.pyplot as plt
import numpy as np

from sklearn.preprocessing import StandardScaler, OneHotEncoder
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import train_test_split, validation_curve
from sklearn.metrics import accuracy_score, precision_score, recall_score, f1_score
from sklearn.metrics import confusion_matrix

from importlib.metadata import version

print(f"Versión de Pandas: {version('pandas')}")
print(f"Versión de Matplotlib: {version('matplotlib')}")
print(f"Versión de Scikit-learn: {version('scikit-learn')}")
print(f"Versión de Numpy: {version('numpy')}")
```

## 1. Carga de datos

Con las librerías importadas, realizaremos la carga del conjunto de datos:

### Ejercicio 1.1.

Utiliza Pandas para importar el archivo que contiene el conjunto de datos de desgaste de empleados.

* La ruta del archivo .csv es: `./data/EmployeeAttrition.csv` y ya se encuentra en el entorno de Coursera, solo debes importarlo.
* La variable resultante debe tener el nombre `data_raw`, que representa el conjunto de datos sin modificar.

<!-- grade_id: E11, points: 0, grade: false, solution: true -->
```python
ruta = './data/Employee-Attrition.csv'
# your code here

```

<!-- grade_id: E11-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# El resultado existe
# El resultado es un DataFrame
# El resultado tiene las dimensiones correctas
#--------------------------------------

# Se verifica que la variable exista
assert data_raw is not None, "Asegúrate de definir la variable /`data_raw/` igualándola a una función de Pandas para leer el archivo."

# Se verifica que sea un DataFrame
assert isinstance(data_raw, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable
assert data_raw.shape == (1470,32), "¿Verificaste que la ruta del archivo CSV y el nombre de la variable son correctos?"
print("¡Los datos tienen las dimensiones correctas!")
```

En este punto es conveniente observar la estructura de los datos. ¿Cuántas variables categóricas tiene este conjunto de datos?

## 2. Preparación de datos

Primero vamos a definir la variable `data` para almacenar un conjunto de datos modificado:

```python
data = data_raw.copy()
```

A continuación, verificaremos si hay valores faltantes:

```python
data.isna().sum()
```

Y verificaremos si hay filas duplicadas:

```python
data.duplicated().sum()
```

Como puedes observar, no tenemos que hacer una limpieza previa, por lo que seguiremos con la preparación de los datos.

### Ejercicio 2.1.

Ahora debes dividir el conjunto de datos en entrenamiento y pruebas. Usando el 80% de los datos para entrenar el modelo y el 20% restante para probarlo, utiliza `scikit-learn` para separar el conjunto de datos en dos.

* Guarda tu respuesta en dos variables: `train` y `test`. (**Ejemplo: `train, test = <<Función>>`**)
* Utiliza el parámetro `random_state=0`. Esto hará que la partición sea siempre la misma.
* Encontrarás la línea `train.head()`. Esta línea se usa para que puedas visualizar el resultado del conjunto de entrenamiento. Déjala al final de la celda y no la modifiques.

<!-- grade_id: E22, points: 0, grade: false, solution: true -->
```python
# your code here

train.head()
```

<!-- grade_id: E22-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables "train" y "test" existen
# Las variables "train" y "test" son un DataFrame
# Las variables tienen las dimensiones correctas
#--------------------------------------

# Se verifica que "train" y "test" están definidas
assert train is not None, "Asegúrate de definir la variable \'train\' con el nombre correcto."
assert test is not None, "Asegúrate de definir la variable \'test\' con el nombre correcto."

# Se verifica que "train" y "test" sean un DataFrame
assert isinstance(train, pd.DataFrame), "La variable \'train\' debe ser un DataFrame."
assert isinstance(test, pd.DataFrame), "La variable \'test\' debe ser un DataFrame."

# Se evalúan las dimensiones de las variables
assert train.shape == (1176,32), "Verifica que estés utilizando el 80% de los datos para el conjunto de entrenamiento."
assert test.shape == (294,32), "Verifica que estés utilizando el 20% de los datos para el conjunto de pruebas."
print("¡Los conjuntos de entrenamiento y pruebas tienen las dimensiones correctas!")
```

### Ejercicio 2.2.

Ahora debes aislar la variable objetivo, `Attrition`, de las variables independientes. Utiliza Pandas para crear dos variables, `x_train` y `y_train`, que almacenarán las variables independientes y la variable objetivo, respectivamente.

* Crea una variable con nombre `x_train` y asígnale la operación necesaria para almacenar solo las variables independientes del conjunto de entrenamiento. (**Ejemplo: `x_train = train.<<Función>>`**)
* Crea una variable con nombre `y_train` y asígnale la operación necesaria para almacenar la variable objetivo del conjunto de entrenamiento. (**Ejemplo: `y_train = <<Consulta>>`**)

<!-- grade_id: E23, points: 0, grade: false, solution: true -->
```python
# Tu respuesta deben ser dos líneas consecutivas:
#    x_train = train.<<Función>>
#    y_train = <<Consulta>>
# your code here

```

<!-- grade_id: E23-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables "x_train" y "y_train" existen
# La variable "x_train" es un DataFrame
# La variable "y_train" es una Serie de Pandas
# Las variables tienen las dimensiones correctas
#--------------------------------------

# Se verifica que las variables están definidas
assert x_train is not None, "Asegúrate de definir la variable \'x_train\' correctamente."
assert y_train is not None, "Asegúrate de definir la variable \'y_train\' correctamente."

# Se verifica que "x_train" sea un DataFrame
assert isinstance(x_train, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se verifica que "y_train" sea una Serie
assert isinstance(y_train, pd.Series), "El resultado debe ser una Serie de Pandas."

# Se evalúan las dimensiones de las variables
assert x_train.shape == (1176,31), "\'x_train\' debe tener el mismo número de filas pero una columna menos que \'train\'."
assert y_train.shape == (1176,), "\'y_train\' solamente contiene una columna."
print("¡Los conjuntos \'x_train\' y \'y_train\' tienen las dimensiones correctas!")
```

### Ejercicio 2.3.

Ahora vas a realizar la transformación de las variables categóricas. Utiliza el encoder definido junto con las variables categóricas y numéricas para transformar todas las variables categóricas en variables numéricas que puedan ser interpretadas por nuestro modelo.


Aquí definimos las variables categóricas y numéricas 

```python
vars_cat= list(x_train.select_dtypes(include=["object",bool], exclude=[int,float]).columns)
vars_num=list(set(x_train.columns)-set(vars_cat))
print(f"variables categóricas: {vars_cat}")
print(f"variables numéricas: {vars_num}")
```

En este paso definimos el OneHotEncoder

```python
encoder = OneHotEncoder(sparse_output=False).set_output(transform = "pandas")
```

En esta parte separamos las variables categóricas de las numéricas

```python
x_train_cat = x_train[vars_cat]
x_train_num = x_train[vars_num]
```

* Utiliza el encoder para **ajustarlo** a las variables categóricas existentes (Ejemplo: **encoder.<<Función>>**)
* Utiliza el encoder para **transformar** las variables categóricas a numéricas (Ejemplo: **x_train_cat = encoder.<<Función>>**)
* Combina nuevamente las variables categóricas y numéricas en un único dataset y guárdalo en la variable `x_train`. (Ejemplo: **pd.concat(<<Parámetros>>)**)
* Encontrarás la línea `x_train.head()` al final de la celda. Esta línea se usa para que puedas visualizar el resultado de la transformación, por lo que debes dejarla al final y no debes modificarla.

<!-- grade_id: E21, points: 0, grade: false, solution: true -->
```python
# your code here

x_train.head()
```

<!-- grade_id: E21-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable 'x_train' existe
# La variable 'x_train' es un DataFrame
# La variable 'x_train' ha aumentado su número de columnas
#--------------------------------------

# Se verifica que la variable exista
assert x_train is not None, "Asegúrate de definir la variable /`data/` igualándola a una función de Pandas."

# Se verifica que sea un DataFrame
assert isinstance(x_train, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable data
assert x_train.shape == (1176,52), "Recuerda que la codificación de variables categóricas resulta en un aumento del número de columnas."
print("¡Los datos tienen las dimensiones correctas!")
```

### Ejercicio 2.4.

El siguiente paso es realizar la estandarización de nuestros datos.

Adicionalmente, definiremos un objeto de la clase `StandardScaler()`, que utilizarás para realizar la estandarización, junto con la configuracion .set_output(transform = "pandas") para mantener las columnas nombradas.:

```python
scaler = StandardScaler().set_output(transform = "pandas")
```

Con las variables definidas, realiza la estandarización de los datos. Modifica el conjunto de variables independientes, `x_train`, con las operaciones necesarias para estandarizarlo. Particularmente, deberás seguir estos pasos:

* Utiliza la variable `scaler` para transformar los datos de `x_train`. Asigna el resultado a la misma variable sobreescribiendo su contenido. (**Ejemplo: `x_train = scaler.<<Función>>`**)
* Reconstruye el DataFrame de variables independientes utilizando la información almacenada en `columns`. Asigna el resultado a `x_train`. (**Ejemplo: `x_train = <<Función>>`**)

<!-- grade_id: E24, points: 0, grade: false, solution: true -->
```python
# your code here

```

<!-- grade_id: E24-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "x_train" es un DataFrame
# La variable tiene las dimensiones correctas
#--------------------------------------

# Se verifica que "x_train" sea un DataFrame
assert isinstance(x_train, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable
assert x_train.shape == (1176,52), "\'x_train\' debe mantener las mismas dimensiones, solo se modifican sus valores."
print("¡\'x_train\' es un DataFrame con las dimensiones correctas!")
```

## 3. Graficación de una curva de validación para el hiperparámetro K

Con el conjunto de entrenamiento preparado, realizaremos una búsqueda del hiperparámetro _k_ mediante una curva de validación. Primero utilizaremos las funciones `list` y `range` para definir nuestro espacio de búsqueda, almacenando el resultado en la variable `k_values`:

```python
k_values = list(range(1, 21))
k_values
```

### Ejercicio 3.1.

Con el espacio de búsqueda, vamos a obtener los valores de desempeño para graficar la curva de validación.

* Define dos variables con los nombres `train_scores` y `val_scores`, asignando la función necesaria para generar los valores de desempeño de una curva de validación (**Ejemplo: `train_scores, val_scores = <<Función>>`**). Utiliza los siguientes parámetros:
    * `estimator = KNeighborsClassifier()`: el algoritmo.
    * `X = x_train`: el conjunto de variables dependientes.
    * `y = y_train`: la variable objetivo.
    * `param_name = 'n_neighbors'`: el hiperparámetro.
    * `param_range = k_values`: el espacio de búsqueda.
    * `scoring = 'accuracy'`: la métrica de desempeño a utilizar.
    * `cv = 10`: los 'folds' para la validación cruzada.

<!-- grade_id: E31, points: 0, grade: false, solution: true -->
```python
# your code here

```

<!-- grade_id: E31-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "train_scores" existe
# La variable "val_scores" existe
# La variable "train_scores" es un arreglo de numpy
# La variable "val_scores" es un arreglo de numpy
# Las variables tienen las dimensiones correctas
#--------------------------------------

# Se verifica que "train_scores" está definida
assert train_scores is not None, "Asegúrate de definir la variable \'train_scores\' correctamente."

# Se verifica que "val_scores" está definida
assert val_scores is not None, "Asegúrate de definir la variable \'val_scores\' correctamente."

# Se verifica que "train_scores" sea un arreglo de numpy
assert isinstance(train_scores, np.ndarray), "El resultado debe ser un arreglo de numpy."

# Se verifica que "val_scores" sea un arreglo de numpy
assert isinstance(val_scores, np.ndarray), "El resultado debe ser un arreglo de numpy."

# Se evalúan las dimensiones de las variables
assert train_scores.shape == (20,10), "\'train_scores\' debe tener dimensiones (20,10)."
assert val_scores.shape == (20,10), "\'val_scores\' debe tener dimensiones (20,10)."
print("¡Los valores de desempeño para la curva de validación se generaron correctamente!")
```

Para generar la curva de validación, vamos a utilizar las funciones `mean` y `std` de `numpy` para obtener la media y desviación estándar en los valores para los conjuntos de entrenamiento y validación:

```python
train_mean = np.mean(train_scores, axis=1)
train_std = np.std(train_scores, axis=1)
val_mean = np.mean(val_scores, axis=1)
val_std = np.std(val_scores, axis=1)
```

A continuación graficaremos la curva de validación con los valores que acabas de obtener. El resultado debe ser similar a la siguiente gráfica:
![Gráfica E3.1](./img/31.png "Gráfica resultante para E3.1.")

```python
# Inicializamos la figura
plt.figure(figsize=(7,7))

# Creamos la curva de validación para el conjunto de entrenamiento
plt.plot(k_values, train_mean, color='blue', marker='o', markersize=5, label='Exactitud en entrenamiento')
plt.fill_between(k_values, train_mean+train_std, train_mean-train_std, alpha=0.15, color='blue')

# Creamos la curva de validación para el conjunto de validación
plt.plot(k_values, val_mean, color='green', linestyle='--', marker='s', markersize=5, label='Exactitud en validación')
plt.fill_between(k_values, val_mean+val_std, val_mean-val_std, alpha=0.15, color='green')

# Agregamos una grilla y etiquetas para los ejes
plt.grid()
plt.legend(loc='lower right')
plt.xlabel('k')
plt.ylabel('Exactitud')
plt.show()
```

### Ejercicio 3.2.

Para obtener el mejor valor de _k_, utilizaremos como criterio la mejor exactitud en el conjunto de validación. Utilizando `numpy`, obtén el valor del hiperparámetro en dos pasos:

* Define una variable con el nombre `mejor_indice` y asígnale la función de `numpy` necesaria para obtener el índice del mayor elemento en `val_mean` (**Ejemplo: `mejor_indice = np.<<Función>>`**).
* Define una variable con el nombre `mejor_k` y asígnale el elemento obtenido al acceder al índice `mejor_indice` del arreglo de valores de _k_ `k_values` (**Ejemplo: `mejor_k = k_values[<<Índice>>]`**).

<!-- grade_id: E32, points: 0, grade: false, solution: true -->
```python
# your code here

mejor_k
```

<!-- grade_id: E32-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "mejor_indice" existe
# La variable "mejor_k" existe
# El valor de "mejor_k" es correcto
#--------------------------------------

# La variable "mejor_indice" está definida
assert mejor_indice is not None, "Asegúrate de definir la variable \'mejor_indice\'."

# La variable "mejor_k" está definida
assert mejor_k is not None, "Asegúrate de definir la variable \'mejor_k\'."

print("¡El valor de K es correcto!")
```

Ahora entrenaremos el modelo de KNN con el mejor valor de _k_:

```python
mejor_modelo = KNeighborsClassifier(n_neighbors=mejor_k).fit(x_train, y_train)
```

## 4. Evaluación del modelo

Por último, evaluarás el modelo entrenado utilizando el conjunto de pruebas.

### Ejercicio 4.0.

Separa las variables independientes y la variable objetivo en el conjunto de pruebas. Para ello, utiliza Pandas para crear dos variables, `x_test` y `y_test`, que almacenarán las variables independientes y la variable objetivo, respectivamente.

* Crea una variable con nombre `x_test` y asígnale la operación necesaria para almacenar solo las variables independientes del conjunto de pruebas. (**Ejemplo: `x_test = test.<<Función>>`**)
* Crea una variable con nombre `y_test` y asígnale la operación necesaria para almacenar la variable objetivo del conjunto de pruebas. (**Ejemplo: `y_test = <<Consulta>>`**)

<!-- grade_id: E40, points: 0, grade: false, solution: true -->
```python
# Tu respuesta deben ser dos líneas consecutivas:
#    x_test = test.<<Función>>
#    y_test = <<Consulta>>
# your code here

```

<!-- grade_id: E40-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables existen
# La variable "x_test" es un DataFrame
# La variable "y_test" es una Serie de Pandas
# Las variables tienen las dimensiones correctas
#--------------------------------------

# Se verifica que las variables están definidas
assert x_test is not None, "Asegúrate de definir la variable con el nombre correcto."
assert y_test is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "x_test" sea un DataFrame
assert isinstance(x_test, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se verifica que "y_test" sea una Serie
assert isinstance(y_test, pd.Series), "El resultado debe ser una Serie de Pandas."

# Se evalúan las dimensiones de las variables
assert x_test.shape == (294,31), "\'x_test\' debe tener el mismo número de filas pero una columna menos que \'test\'."
assert y_test.shape == (294,), "\'y_test\' solamente contiene una columna."
print("¡Los conjuntos \'x_test\' y \'y_test\' tienen las dimensiones correctas!")
```

### Ejercicio 4.1.

Realiza la transformación de las variables categóricas. Utiliza el encoder definido junto con las variables categóricas y numéricas para transformar todas las variables categóricas en variables numéricas que puedan ser interpretadas por nuestro modelo.

1. Separa los datos de test en categóricos (con las columnas de la variable `vars_cat`) y numéricos (con las columnas de la variable `vars_num`) usando pandas
2. Codifica los datos categóricos usando el encoder (Ejemplo: `categoricos_encoded = encoder.<<Función>>`)
3. Junta nuevamente los datos numéricos y categóricos en un único conjunto de test, y guárdalos en `x_test`, sobreescribiendo su contenido (Ejemplo: `x_test = pd.<<Función>>`)

<!-- grade_id: E41, points: 0, grade: false, solution: true -->
```python
# your code here

x_test.head()
```

<!-- grade_id: E41-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "x_test" es un DataFrame
# La variable tiene las dimensiones correctas
#--------------------------------------

# Se verifica que "x_test" sea un DataFrame
assert isinstance(x_test, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable
assert x_test.shape == (294,52), "\'x_test\' debe mantener las mismas dimensiones, solo se modifican sus valores."
print("¡\'x_test\' es un DataFrame con las dimensiones correctas!")
```

### Ejercicio 4.2.

Realiza la estandarización de los datos del conjunto de pruebas. Modifica el conjunto de variables independientes, `x_test`, con las operaciones necesarias para estandarizarlo utilizando **solamente la información del conjunto de entrenamiento**. Particularmente, deberás seguir estos pasos:

* Utiliza la variable `scaler` para transformar los datos de `x_test`. Asigna el resultado a la misma variable, sobreescribiendo su contenido. (**Ejemplo: `x_test = scaler.<<Función>>`**)

<!-- grade_id: E42, points: 0, grade: false, solution: true -->
```python
# your code here

```

<!-- grade_id: E42-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "x_test" es un DataFrame
# La variable tiene las dimensiones correctas
#--------------------------------------

# Se verifica que "x_test" sea un DataFrame
assert isinstance(x_test, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable
assert x_test.shape == (294,52), "\'x_test\' debe mantener las mismas dimensiones, solo se modifican sus valores."
print("¡\'x_test\' es un DataFrame con las dimensiones correctas!")
```

### Ejercicio 4.3.

Con el conjunto de pruebas preparado, realiza predicciones con el fin de compararlas con los valores reales almacenados en `y_test`.

* Utiliza la variable `mejor_modelo` para realizar las predicciones sobre el mejor modelo. Asigna el resultado a una variable con nombre `y_pred` (**Ejemplo: `y_pred = mejor_modelo.<<Función>>`**).

<!-- grade_id: E43, points: 0, grade: false, solution: true -->
```python
# your code here

```

<!-- grade_id: E43-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "y_pred" existe
# La variable "y_pred" es un arreglo
# La variable "y_pred" tiene las dimensiones correctas
#--------------------------------------

# Se verifica que "y_pred" está definida
assert y_pred is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "y_pred" sea un arreglo
assert isinstance(y_pred, np.ndarray), "El resultado debe ser un arreglo de Numpy."

# Se evalúan las dimensiones de "y_pred"
assert y_pred.shape == (294,), "\'y_pred\' debe tener el mismo número de filas (predicciones) que \'x_test\' y \'y_test\'."
print("¡\'y_pred\' es un arreglo con las dimensiones correctas!")
```

### Ejercicio 4.4.

Utiliza el conjunto de predicciones `y_pred` y el conjunto de valores reales `y_test` para obtener la matriz de confusión del mejor modelo.

* Haz un llamado a la función que retorna la matriz de confusión como un arreglo. Asigna el resultado a una nueva variable con el nombre `p44` (**Ejemplo: `p44 = <<Función>>`**).
* Encontrarás la línea `p44` al final de la celda. Esta línea se usa para que puedas visualizar la matriz resultante, por lo que no la debes modificar.

<!-- grade_id: E44, points: 0, grade: false, solution: true -->
```python
# your code here

p44
```

<!-- grade_id: E44-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "p44" existe
# La variable "p44" es un arreglo de numpy
# La variable "p44" tiene las dimensiones correctas
#--------------------------------------

# Se verifica que "p44" está definida
assert p44 is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "p44" sea un arreglo de numpy
assert isinstance(p44, np.ndarray), "El resultado debe ser un arreglo de Numpy."

# Se evalúan las dimensiones de "p44"
assert p44.shape == (2,2), "\'p44\' debe tener dos filas y dos columnas, equivalente al número de clases del problema."
print("¡\'p44\' es un arreglo con las dimensiones correctas!")
```

### Ejercicio 4.5.

Finalmente, obtén cuatro métricas de rendimiento utilizando las predicciones del modelo.

#### Ejercicio 4.5.1.

Utiliza `scikit-learn` para obtener la exactitud (accuracy).

* Define una variable con el nombre `accuracy` y asígnale la función necesaria para obtener la exactitud (`accuracy = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E451, points: 0, grade: false, solution: true -->
```python
# your code here

print("Exactitud: ", accuracy)
```

<!-- grade_id: E451-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "accuracy" existe
# La variable "accuracy" es un número
# La variable "accuracy" tiene un valor válido
#--------------------------------------

# Se verifica que "accuracy" está definida
assert accuracy is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "accuracy" sea un número
assert isinstance(accuracy, np.float64) or isinstance(accuracy, float), "El resultado debe ser un número decimal."

# Se verifica que "accuracy" tiene un valor válido
assert accuracy >= 0 and accuracy <= 1, "\'accuracy\' debe tener un valor entre 0 y 1."
print("¡\'accuracy\' es un número válido!")
```

#### Ejercicio 4.5.2.

Utiliza `scikit-learn` para obtener la sensibilidad (recall).

* Define una variable con el nombre `recall` y asígnale la función necesaria para obtener la sensibilidad (`recall = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E452, points: 0, grade: false, solution: true -->
```python
# your code here

print("Sensibilidad: ", recall)
```

<!-- grade_id: E452-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "recall" existe
# La variable "recall" es un número
# La variable "recall" tiene un valor válido
#--------------------------------------

# Se verifica que "recall" está definida
assert recall is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "recall" sea un número
assert isinstance(recall, np.float64) or isinstance(recall, float), "El resultado debe ser un número decimal."

# Se verifica que "recall" tiene un valor válido
assert recall >= 0 and recall <= 1, "\'recall\' debe tener un valor entre 0 y 1."
print("¡\'recall\' es un número válido!")
```

#### Ejercicio 4.5.3.

Utiliza `scikit-learn` para obtener la precisión.

* Define una variable con el nombre `precision` y asígnale la función necesaria para obtener la precisión (`precision = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E453, points: 0, grade: false, solution: true -->
```python
# your code here

print('Precisión: ', precision)
```

<!-- grade_id: E453-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "precision" existe
# La variable "precision" es un número
# La variable "precision" tiene un valor válido
#--------------------------------------

# Se verifica que "precision" está definida
assert precision is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "precision" sea un número
assert isinstance(precision, np.float64) or isinstance(precision, float), "El resultado debe ser un número decimal."

# Se verifica que "precision" tiene un valor válido
assert precision >= 0 and precision <= 1, "\'precision\' debe tener un valor entre 0 y 1."
print("¡\'precision\' es un número válido!")
```

#### Ejercicio 4.5.4.

Utiliza `scikit-learn` para obtener el F1-Score.

* Define una variable con el nombre `f1` y asígnale la función necesaria para obtener el F1-Score (`f1 = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E454, points: 0, grade: false, solution: true -->
```python
# your code here

print('F1-Score: ', f1)
```

<!-- grade_id: E454-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "f1" existe
# La variable "f1" es un número
# La variable "f1" tiene un valor válido
#--------------------------------------

# Se verifica que "f1" está definida
assert f1 is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "f1_score" sea un número
assert isinstance(f1, np.float64) or isinstance(f1, float), "El resultado debe ser un número decimal."

# Se verifica que "f1" tiene un valor válido
assert f1 >= 0 and f1 <= 1, "\'f1_score\' debe tener un valor entre 0 y 1."
print("¡\'f1\' es un número válido!")
```

## Cierre

Al realizar los ejercicios de este taller, has reforzado tus capacidades para realizar una búsqueda exhaustiva de hiperparámetros mediante una curva de validación. Además, lograste evaluar el mejor modelo de KNN resultante utilizando la matriz de confusión y cuatro métricas de rendimiento. 

---
Creado por: Nicolás Díaz

Última edición: Camilo Rozo

Revisado por: Haydemar Nuñez

Versión: Febrero 2025

Universidad de los Andes