# Un problema de regresión logística

En este taller estudiarás conceptos mostrados en el tutorial "Solución de problemas con regresión logística", utilizando un nuevo conjunto de datos correspondiente al desempeño y desgaste de empleados. Particularmente, realizarás los siguientes procesos:

1. Cargar un conjunto de datos.
2. Preparar los datos para el modelado.
3. Realizar la búsqueda de hiperparámetros para la regresión logística.
4. Evaluar el mejor modelo resultante.
5. Revisar indicios de bajoajuste o sobreajuste


El problema a resolver es el siguiente: dadas algunas características en áreas como la educación, trabajos previos, salario, entre otras, queremos clasificar a un empleado como agotado (con bajo desempeño) o no. Antes de iniciar, vamos a importar las librerías necesarias:

```python
import pandas as pd
import matplotlib.pyplot as plt
import numpy as np

from sklearn.preprocessing import StandardScaler, OneHotEncoder
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import train_test_split, KFold, GridSearchCV
from sklearn.metrics import accuracy_score, precision_score, recall_score, f1_score
from sklearn.metrics import confusion_matrix, classification_report
from sklearn.metrics import ConfusionMatrixDisplay

from importlib.metadata import version
```

```python
print(f"Versión de Pandas: {version('pandas')}")
print(f"Versión de Matplotlib: {version('matplotlib')}")
print(f"Versión de Scikit-learn: {version('scikit-learn')}")
print(f"Versión de Numpy: {version('numpy')}")
```

Output:
```
Versión de Pandas: 2.2.3
Versión de Matplotlib: 3.10.0
Versión de Scikit-learn: 1.6.1
Versión de Numpy: 1.25.2
```

## 1. Carga de datos

Con las librerías importadas, realizaremos la carga del conjunto de datos:

### Ejercicio 1.1.

Utiliza Pandas para importar el archivo que contiene el conjunto de datos de desgaste de empleados.

* La ruta del archivo .csv es: `./data/EmployeeAttrition.csv`, y ya se encuentra en el entorno de Coursera, solo debes importarlo.
* La variable resultante debe tener el nombre `data_raw`, que representa el conjunto de datos sin modificar.

<!-- grade_id: E11, points: 0, grade: false, solution: true -->
```python
ruta = './data/Employee-Attrition.csv'
# your code here
data_raw = pd.read_csv(ruta, sep=',')
```

<!-- grade_id: E11-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# El resultado existe
# El resultado es un DataFrame
# El resultado tiene las dimensiones correctas
#--------------------------------------

# Se verifica que la variable exista
assert data_raw is not None, "Asegúrate de definir la variable /`data_raw/` igualándola a una función de Pandas para leer el archivo."
assert ruta is not None, "Asegúrate de definir la ruta de los datos correctamente en una variable `ruta`." 
# Se verifica que sea un DataFrame
assert isinstance(data_raw, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable
assert data_raw.shape == (1470,32), "¿Verificaste que la ruta del archivo CSV y el nombre de la variable son correctos?"
print("¡Los datos tienen las dimensiones correctas!")
```

Output:
```
¡Los datos tienen las dimensiones correctas!
```

En este punto es conveniente observar la estructura de los datos. ¿Cuántas variables categóricas tiene este conjunto de datos?

```python
data_raw.head()
```

Output:
```
Age  Attrition     BusinessTravel  DailyRate              Department  \
0   41          1      Travel_Rarely       1102                   Sales   
1   49          0  Travel_Frequently        279  Research & Development   
2   37          1      Travel_Rarely       1373  Research & Development   
3   33          0  Travel_Frequently       1392  Research & Development   
4   27          0      Travel_Rarely        591  Research & Development   

   DistanceFromHome  Education EducationField  EmployeeNumber  \
0                 1          2  Life Sciences               1   
1                 8          1  Life Sciences               2   
2                 2          2          Other               4   
3                 3          4  Life Sciences               5   
4                 2          1        Medical               7   

   EnvironmentSatisfaction  ... PerformanceRating  RelationshipSatisfaction  \
0                        2  ...                 3                         1   
1  
```

## 2. Preparación de datos

Primero vamos a definir la variable `data` para almacenar un conjunto de datos modificado:

```python
data = data_raw.copy()
```

### Ejercicio 2.1.

Primero debes dividir el conjunto de datos en entrenamiento y pruebas. Usando el 80% de los datos para entrenar el modelo y el 20% restante para probarlo, utiliza `scikit-learn` para separar el conjunto de datos en dos.

* Guarda tu respuesta en dos variables: `train` y `test`. (**Ejemplo: `train, test = <<Función>>`**)
* Utiliza el parámetro `random_state=0`. Esto hará que la partición sea siempre la misma.
* Luego de separar los datos, reinicia el índice de train (**Ejemplo: `train = train.<<Función>>`**) y de test por separado, si lo haces correctamente, los números de filas deberían iniciar en 0 e ir incrementando en vez de ser aleatorios
* Encontrarás la línea `train.head()`. Esta línea se usa para que puedas visualizar el resultado del conjunto de entrenamiento. Déjala al final de la celda y no la modifiques.

<!-- grade_id: E22, points: 0, grade: false, solution: true -->
```python
# your code here
train, test = train_test_split(data, test_size=0.2, random_state=0)
train = train.reset_index(drop=True)
train.head()
```

Output:
```
Age  Attrition     BusinessTravel  DailyRate              Department  \
0   58          0      Travel_Rarely        605                   Sales   
1   45          0      Travel_Rarely        950  Research & Development   
2   40          0      Travel_Rarely        300                   Sales   
3   36          0         Non-Travel       1434                   Sales   
4   25          1  Travel_Frequently        599                   Sales   

   DistanceFromHome  Education    EducationField  EmployeeNumber  \
0                21          3     Life Sciences            1938   
1                28          3  Technical Degree            1546   
2                26          3         Marketing            1066   
3                 8          4     Life Sciences             789   
4                24          1     Life Sciences            1273   

   EnvironmentSatisfaction  ... PerformanceRating  RelationshipSatisfaction  \
0                        4  ...                 3               
```

<!-- grade_id: E22-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables "train" y "test" existen
# Las variables "train" y "test" son un DataFrame
# Las variables tienen las dimensiones correctas
#--------------------------------------

# Se verifica que "train" y "test" están definidas
assert train is not None, "Asegúrate de definir la variable \'train\' con el nombre correcto."
assert test is not None, "Asegúrate de definir la variable \'test\' con el nombre correcto."

# Se verifica que "train" y "test" sean un DataFrame
assert isinstance(train, pd.DataFrame), "La variable \'train\' debe ser un DataFrame."
assert isinstance(test, pd.DataFrame), "La variable \'test\' debe ser un DataFrame."

assert "index" not in train.columns, "Asegúrate de eliminar el índice anterior con drop=True"
# Se evalúan las dimensiones de las variables
assert train.shape == (1176,32), "Verifica que estés utilizando el 80% de los datos para el conjunto de entrenamiento."
assert test.shape == (294,32), "Verifica que estés utilizando el 20% de los datos para el conjunto de pruebas."
print("¡Los conjuntos de entrenamiento y pruebas tienen las dimensiones correctas!")
```

Output:
```
¡Los conjuntos de entrenamiento y pruebas tienen las dimensiones correctas!
```

A continuación, verificaremos si hay valores faltantes:

```python
train.isna().sum()
```

Output:
```
Age                         0
Attrition                   0
BusinessTravel              0
DailyRate                   0
Department                  0
DistanceFromHome            0
Education                   0
EducationField              0
EmployeeNumber              0
EnvironmentSatisfaction     0
Gender                      0
HourlyRate                  0
JobInvolvement              0
JobLevel                    0
JobRole                     0
JobSatisfaction             0
MaritalStatus               0
MonthlyIncome               0
MonthlyRate                 0
NumCompaniesWorked          0
OverTime                    0
PercentSalaryHike           0
PerformanceRating           0
RelationshipSatisfaction    0
StockOptionLevel            0
TotalWorkingYears           0
TrainingTimesLastYear       0
WorkLifeBalance             0
YearsAtCompany              0
YearsInCurrentRole          0
YearsSinceLastPromotion     0
YearsWithCurrManager        0
dtype: int64
```

Y verificaremos si hay filas duplicadas:

```python
data.duplicated().sum()
```

Output:
```
0
```

Como puedes observar, no tenemos que hacer una limpieza previa, por lo que seguiremos con la preparación de los datos.

### Ejercicio 2.2.

Ahora debes aislar la variable objetivo, `Attrition`, de las variables independientes. Utiliza Pandas para crear dos variables, `x_train` y `y_train`, que almacenarán las variables independientes y la variable objetivo, respectivamente.

* Crea una variable con nombre `x_train` y asígnale la operación necesaria para almacenar solo las variables independientes del conjunto de entrenamiento. (**Ejemplo: `x_train = train.<<Función>>`**)
* Crea una variable con nombre `y_train` y asígnale la operación necesaria para almacenar la variable objetivo del conjunto de entrenamiento. (**Ejemplo: `y_train = <<Consulta>>`**)

<!-- grade_id: E23, points: 0, grade: false, solution: true -->
```python
# Tu respuesta deben ser dos líneas consecutivas:
#    x_train = train.<<Función>>
#    y_train = <<Consulta>>
# your code here
x_train = train.drop(['Attrition'],axis="columns")
y_train = train['Attrition']
```

<!-- grade_id: E23-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables "x_train" y "y_train" existen
# La variable "x_train" es un DataFrame
# La variable "y_train" es una Serie de Pandas
# Las variables tienen las dimensiones correctas
#--------------------------------------

# Se verifica que las variables están definidas
assert x_train is not None, "Asegúrate de definir la variable \'x_train\' correctamente."
assert y_train is not None, "Asegúrate de definir la variable \'y_train\' correctamente."

# Se verifica que "x_train" sea un DataFrame
assert isinstance(x_train, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se verifica que "y_train" sea una Serie
assert isinstance(y_train, pd.Series), "El resultado debe ser una Serie de Pandas."

# Se evalúan las dimensiones de las variables
assert x_train.shape == (1176,31), "\'x_train\' debe tener el mismo número de filas pero una columna menos que \'train\'."
assert y_train.shape == (1176,), "\'y_train\' solamente contiene una columna."
print("¡Los conjuntos \'x_train\' y \'y_train\' tienen las dimensiones correctas!")
```

Output:
```
¡Los conjuntos 'x_train' y 'y_train' tienen las dimensiones correctas!
```

### Ejercicio 2.3.

Primero vas a realizar la transformación de las variables categóricas. Para esto, vamos a usar un `OneHotEncoder` de scikit-learn

Guardaremos los nombres de las variables categóricas y numéricas:

```python
nombres_variables_categoricas = ["BusinessTravel","Department","EducationField","Gender","JobRole","MaritalStatus","OverTime"]
nombres_variables_numericas = list(set(x_train.columns) - set(nombres_variables_categoricas))
print(f"Variables categóricas: {nombres_variables_categoricas}\n")
print(f"Variables numéricas: {nombres_variables_numericas}")
```

Output:
```
Variables categóricas: ['BusinessTravel', 'Department', 'EducationField', 'Gender', 'JobRole', 'MaritalStatus', 'OverTime']

Variables numéricas: ['NumCompaniesWorked', 'TrainingTimesLastYear', 'DailyRate', 'JobLevel', 'TotalWorkingYears', 'YearsInCurrentRole', 'StockOptionLevel', 'YearsSinceLastPromotion', 'JobSatisfaction', 'PercentSalaryHike', 'PerformanceRating', 'WorkLifeBalance', 'DistanceFromHome', 'HourlyRate', 'RelationshipSatisfaction', 'MonthlyRate', 'Age', 'YearsAtCompany', 'MonthlyIncome', 'EmployeeNumber', 'EnvironmentSatisfaction', 'YearsWithCurrManager', 'JobInvolvement', 'Education']
```

Creamos el objeto del encoder

```python
encoder = OneHotEncoder(sparse_output=False)
```

Separamos los datos categóricos de los numéricos

```python
datos_categoricos = x_train[nombres_variables_categoricas]
datos_numericos = x_train[nombres_variables_numericas]
```

Con los datos separados, realiza el one-hot encoding. Modifica el conjunto de variables independientes, `x_train`, con las operaciones necesarias para estandarizarlo. Particularmente, deberás seguir estos pasos:

* Utiliza la variable `encoder` para transformar los datos de `datos_categoricos` (**Ejemplo: `encoder.<<Función>>`**)
* Puedes usar la función `get_feature_names_out()` del encoder para obtener los nombres de las nuevas variables binarias.

<!-- grade_id: cell-4252508dcc3f41cf, points: 0, grade: false, solution: true -->
```python
# your code here
encoder.fit(datos_categoricos)
pd.DataFrame(encoder.get_feature_names_out(),columns=["Variable"])
```

Output:
```
Variable
0           BusinessTravel_Non-Travel
1    BusinessTravel_Travel_Frequently
2        BusinessTravel_Travel_Rarely
3          Department_Human Resources
4   Department_Research & Development
5                    Department_Sales
6      EducationField_Human Resources
7        EducationField_Life Sciences
8            EducationField_Marketing
9              EducationField_Medical
10               EducationField_Other
11    EducationField_Technical Degree
12                      Gender_Female
13                        Gender_Male
14  JobRole_Healthcare Representative
15            JobRole_Human Resources
16      JobRole_Laboratory Technician
17                    JobRole_Manager
18     JobRole_Manufacturing Director
19          JobRole_Research Director
20         JobRole_Research Scientist
21            JobRole_Sales Executive
22       JobRole_Sales Representative
23             MaritalStatus_Divorced
24              MaritalStatus_Married
25               MaritalStatus_Single
26 
```

<!-- grade_id: cell-5c949eb08847070c, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# se ajustó el encoder
# Las variables "nombres_variables_categoricas" y "datos_categoricos" existen
# El encoder contiene las nuevas columnas de los datos categóricos
#--------------------------------------

assert len(encoder.get_feature_names_out()) > 0, "asegúrate de ajustar el encoder a las variables categóricas"
assert nombres_variables_categoricas is not None, "asegúrate de guardar los nombres de las columnas categóricas en ´nombres_variables_categoricas´" 
assert datos_categoricos is not None, "asegúrate de definir la variable ´datos_categoricos´"
print("¡codificaste las variables categóricas correctamente!")
```

Output:
```
¡codificaste las variables categóricas correctamente!
```

Ahora, con el encoder ajustado, transformaremos las variables categóricas a binarias y las volveremos a juntar con el resto de los datos:

* Usa la función adecuada del encoder para transformar las variables categóricas de `datos_categoricos` (**Ejemplo: `encoder.<<Función>>`**) y guarda el resultado en la variable `x_train_cat`.
* Luego, convierte ese resultado en dataframe nuevamente y asegúrate de agregar los nuevos nombres de las columnas, puedes obtenerlos con `encoder.get_feature_names_out()`

<!-- grade_id: E21, points: 0, grade: false, solution: true -->
```python
# your code here
x_train_cat = pd.DataFrame(encoder.transform(datos_categoricos),columns=encoder.get_feature_names_out())

x_train_cat.head()
```

Output:
```
BusinessTravel_Non-Travel  BusinessTravel_Travel_Frequently  \
0                        0.0                               0.0   
1                        0.0                               0.0   
2                        0.0                               0.0   
3                        1.0                               0.0   
4                        0.0                               1.0   

   BusinessTravel_Travel_Rarely  Department_Human Resources  \
0                           1.0                         0.0   
1                           1.0                         0.0   
2                           1.0                         0.0   
3                           0.0                         0.0   
4                           0.0                         0.0   

   Department_Research & Development  Department_Sales  \
0                                0.0               1.0   
1                                1.0               0.0   
2                                0.0               1.
```

<!-- grade_id: E21-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# la variable "x_train_cat" existe
# la variable "x_train_cat" es un dataframe y tiene las columnas debidamente nombradas
# la variable "x_train_cat" se definió correctamente
# la variable "x_train_cat" tiene las dimensiones correctas
# se ajustó el encoder
#--------------------------------------

# Se verifica que la variable exista
assert x_train_cat is not None, "Asegúrate de definir la variable /`x_train_cat/` igualándola a una función del encoder."
assert len(encoder.get_feature_names_out()) > 0, "asegúrate de ajustar el encoder a las variables categóricas"
assert nombres_variables_categoricas is not None, "asegúrate de guardar los nombres de las columnas categóricas en ´nombres_variables_categoricas´" 
assert datos_categoricos is not None, "asegúrate de definir la variable ´datos_categoricos´"

# Se verifica que sea un DataFrame
assert isinstance(x_train_cat, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable
assert x_train_cat.shape == (1176,28), "Recuerda que la codificación de variables categóricas resulta en un aumento del número de columnas."
print("¡Los datos tienen las dimensiones correctas!")
```

Output:
```
¡Los datos tienen las dimensiones correctas!
```

Por último en este paso, con las variables categóricas, debemos juntar nuevamente los datos categóricos con los numéricos.

* las variables categóricas transformadas se encuentran en la variable `x_train_cat`, las variables numéricas se encuentran en `datos_numericos`. Usa una función de pandas para combinar los datos nuevamente en único dataset de entrenamiento (**Ejemplo: `pd.concat(<<Parámetros>>)`**). Guarda el resultado en la variable `x_train`, reemplazando el dataset de entrenamiento original.

<!-- grade_id: cell-510c621bc1a95659, points: 0, grade: false, solution: true -->
```python
x_train
# your code here
x_train = pd.concat([datos_numericos,x_train_cat],axis="columns")

x_train.head()
```

Output:
```
NumCompaniesWorked  TrainingTimesLastYear  DailyRate  JobLevel  \
0                   4                      2        605         4   
1                   4                      3        950         1   
2                   1                      3        300         2   
3                   1                      1       1434         3   
4                   1                      4        599         1   

   TotalWorkingYears  YearsInCurrentRole  StockOptionLevel  \
0                 29                   0                 1   
1                  8                   4                 1   
2                  8                   7                 1   
3                 10                   7                 0   
4                  1                   0                 0   

   YearsSinceLastPromotion  JobSatisfaction  PercentSalaryHike  ...  \
0                        0                4                 13  ...   
1                        0                4                 20  ...   
2 
```

<!-- grade_id: cell-ebc5e507aad447c4, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# la variable "x_train" existe
# la variable "x_train" es un dataframe y tiene las columnas debidamente nombradas
# la variable "x_train" tiene las dimensiones correctas
#--------------------------------------
assert x_train is not None,  "asegúrate de definir la variable ´x_train´"
assert x_train_cat is not None, "asegúrate de definir la variable ´x_train_cat´"
assert datos_numericos is not None, "asegúrate de definir la variable ´datos_numericos´"

# Se evalúan las dimensiones de la variable
assert x_train.shape == (1176,52), "Recuerda que debes combinar nuevamente las columnas numéricas con las categóricas, la cantidad de filas no debería cambiar"

print("¡Los datos de entrenamiento se juntaron correctamente!")

```

Output:
```
¡Los datos de entrenamiento se juntaron correctamente!
```

### Ejercicio 2.4.

El siguiente paso es realizar la estandarización de nuestros datos. Inicialmente, vamos a almacenar los nombres de las columnas del conjunto de entrenamiento en la variable `columns`:

```python
columns = x_train.columns
```

Adicionalmente, definiremos un objeto de la clase `StandardScaler()`, que utilizarás para realizar la estandarización:

```python
scaler = StandardScaler()
```

Con las variables definidas, realiza la estandarización de los datos. Modifica el conjunto de variables independientes, `x_train`, con las operaciones necesarias para estandarizarlo. Particularmente, deberás seguir estos pasos:

* Utiliza la variable `scaler` para transformar los datos de `x_train`. Asigna el resultado a la misma variable sobreescribiendo su contenido. (**Ejemplo: `x_train = scaler.<<Función>>`**)
* Reconstruye el DataFrame de variables independientes utilizando la información almacenada en `columns`. Asigna el resultado a `x_train`. (**Ejemplo: `x_train = <<Función>>`**)

<!-- grade_id: E24, points: 0, grade: false, solution: true -->
```python
# your code here
x_train = scaler.fit_transform(x_train)
x_train = pd.DataFrame(x_train, columns=columns)
```

<!-- grade_id: E24-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "x_train" es un DataFrame
# La variable tiene las dimensiones correctas
#--------------------------------------

# Se verifica que "x_train" sea un DataFrame
assert isinstance(x_train, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable
assert x_train.shape == (1176,52), "\'x_train\' debe mantener las mismas dimensiones, solo se modifican sus valores."
print("¡\'x_train\' es un DataFrame con las dimensiones correctas!")
```

Output:
```
¡'x_train' es un DataFrame con las dimensiones correctas!
```

## 3. Búsqueda de hiperparámetros

Con el conjunto de entrenamiento preparado, realizaremos una búsqueda de los hiperparámetros `C`, `penalty` y `solver`. Para esto utilizaremos un objeto de la clase `KFold`, con el que definiremos 10 subconjuntos sobre el conjunto de entrenamiento:

```python
kfold = KFold(n_splits=10, shuffle=True, random_state=0)
```

Además, definiremos el objeto de la clase `LogisticRegression` para realizar el entrenamiento con regresión logística:

```python
reglog = LogisticRegression()
```

### Ejercicio 3.1.

El siguiente paso es definir el espacio de búsqueda de los hiperparámetros.

* Define una variable con el nombre `param_grid` y asígnale la expresión necesaria para crear un diccionario con dos tuplas (**Ejemplo: `param_grid = <<Expresión>>`**):
    * Llave `C` y valor `valores_C`.
    * Llave `solver` y valor `valores_solver`.

<!-- grade_id: E31, points: 0, grade: false, solution: true -->
```python
valores_C = [0.1, 0.5, 1]
valores_solver = ['newton-cg', 'lbfgs', 'liblinear']
# your code here
param_grid = {'C': valores_C,
              'solver': valores_solver}
```

<!-- grade_id: E31-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "param_grid" existe
# La variable "param_grid" es un diccionario
# La variable tiene la longitud correcta
#--------------------------------------

# Se verifica que "param_grid" está definida
assert param_grid is not None, "Asegúrate de definir la variable \'param_grid\' correctamente."

# Se verifica que "param_grid" sea un diccionario
assert isinstance(param_grid, dict), "El resultado debe ser un diccionario de Python, es decir, no necesitas ninguna función específica de ninguna librería para definir la variable."

# Se evalúa la longitud de la variable
assert len(param_grid) == 2, "\'param_grid\' debe contener dos tuplas."
print("¡\'param_grid\' está definida correctamente!")
```

Output:
```
¡'param_grid' está definida correctamente!
```

### Ejercicio 3.2.

Finalmente, el último paso antes de realizar la búsqueda de hiperparámetros es crear el objeto de tipo `GridSearchCV`. Utiliza las variables `reglog`, `param_grid` y `kfold` para definirlo:

* Define una variable con el nombre `grid` y asígnale la función necesaria para crear un objeto de la clase `GridSearchCV`. (**Ejemplo: `grid = <<Función>>`**)
* Utiliza el parámetro `scoring='accuracy'` para que se seleccione el mejor modelo de acuerdo con los valores de exactitud.
* Utiliza el parámetro `n_jobs=1` para evitar errores al correr tareas en paralelo.

<!-- grade_id: E32, points: 0, grade: false, solution: true -->
```python
# your code here
grid = GridSearchCV(reglog, param_grid, cv=kfold, scoring='accuracy', n_jobs=-1)
```

<!-- grade_id: E32-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "grid" existe
# La variable "grid" es un objeto de la clase GridSearchCV
# La variable "grid" usa la exactitud como método de selección
#--------------------------------------

# Se verifica que "grid" está definida
assert grid is not None, "Asegúrate de definir la variable \'grid\' correctamente."

# Se verifica que "grid" es un objeto de la clase GridSearchCV
assert type(grid) == GridSearchCV, "La variable \'grid\' debe ser un objeto de la clase GridSearchCV."

# Se verifica que "grid" tiene un atributo "scoring" igual a "accuracy"
assert grid.scoring == 'accuracy', "La variable \'grid\' debe usar la exactitud como método de selección."
print("¡\'grid\' está definida correctamente!")
```

Output:
```
¡'grid' está definida correctamente!
```

### Ejercicio 3.3.

A continuación, realiza la búsqueda de hiperparámetros utilizando el conjunto de entrenamiento, compuesto por las variables `x_train` y `y_train`.

* Para este ejercicio no debes asignar tu resultado a ninguna variable. Es decir, solo debes ejecutar una función sobre la variable `grid`, utilizando las variables `x_train` y `y_train` como parámetros. (**Ejemplo: `grid.<<Función>>`**)

<!-- grade_id: E33, points: 0, grade: false, solution: true -->
```python
# your code here
grid.fit(x_train, y_train)
```

Output:
```
GridSearchCV(cv=KFold(n_splits=10, random_state=0, shuffle=True),
             estimator=LogisticRegression(), n_jobs=-1,
             param_grid={'C': [0.1, 0.5, 1],
                         'solver': ['newton-cg', 'lbfgs', 'liblinear']},
             scoring='accuracy')
```

<!-- grade_id: E33-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# El atributo "best_params_" de la variable "grid" existe
# El atributo "best_estimator_" de la variable "grid" existe
#--------------------------------------

# El atributo "best_params_" está definido
assert grid.best_params_ is not None, "Asegúrate de ejecutar la función de entrenamiento para generar un diccionario con los mejores hiperparámetros."

# El atributo "best_estimator_" está definido
assert grid.best_estimator_ is not None, "Asegúrate de ejecutar la función de entrenamiento para generar un modelo de regresión logística."

# Se verifica que "grid" haya generado dos mejores hiperparámetros
assert len(grid.best_params_) == 2, "Al ejecutar una función usando el objeto \'grid\', se deben generar dos tuplas con los hiperparámetros."

# Se verifica que "grid" haya generado un modelo de regresión logística
assert isinstance(grid.best_estimator_, LogisticRegression), "Al ejecutar una función usando el objeto \'grid\', se debe generar un modelo de regresión logística."
print("¡Se ha realizado la búsqueda de hiperparámetros correctamente!")
```

Output:
```
¡Se ha realizado la búsqueda de hiperparámetros correctamente!
```

Ahora obtendremos los mejores valores de los hiperparámetros usando `grid.best_params_`:

```python
print("Mejores parámetros: {}".format(grid.best_params_)) 
```

Output:
```
Mejores parámetros: {'C': 0.1, 'solver': 'newton-cg'}
```

Además, almacenaremos el mejor modelo utilizando `grid.best_estimator_`:

```python
mejor_modelo = grid.best_estimator_
```

## 4. Evaluación del modelo

Por último, evaluarás el modelo entrenado utilizando el conjunto de pruebas.

### Ejercicio 4.0.

Separa las variables independientes y la variable objetivo en el conjunto de pruebas. Para ello, utiliza Pandas para crear dos variables, `x_test` y `y_test`, que almacenarán las variables independientes y la variable objetivo, respectivamente.

* Crea una variable con nombre `x_test` y asígnale la operación necesaria para almacenar solo las variables independientes del conjunto de pruebas. (**Ejemplo: `x_test = test.<<Función>>`**)
* Crea una variable con nombre `y_test` y asígnale la operación necesaria para almacenar la variable objetivo del conjunto de pruebas. (**Ejemplo: `y_test = <<Consulta>>`**)

<!-- grade_id: E41, points: 0, grade: false, solution: true -->
```python
# Tu respuesta deben ser dos líneas consecutivas:
#    x_test = test.<<Función>>
#    y_test = <<Consulta>>
# your code here
x_test = test.drop(['Attrition'],axis="columns")
y_test = test['Attrition']
```

<!-- grade_id: E41-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables existen
# La variable "x_test" es un DataFrame
# La variable "y_test" es una Serie de Pandas
# Las variables tienen las dimensiones correctas
#--------------------------------------

# Se verifica que las variables están definidas
assert x_test is not None, "Asegúrate de definir la variable con el nombre correcto."
assert y_test is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "x_test" sea un DataFrame
assert isinstance(x_test, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se verifica que "y_test" sea una Serie
assert isinstance(y_test, pd.Series), "El resultado debe ser una Serie de Pandas."

# Se evalúan las dimensiones de las variables
assert x_test.shape == (294,31), "\'x_test\' debe tener el mismo número de filas pero una columna menos que \'test\'."
assert y_test.shape == (294,), "\'y_test\' solamente contiene una columna."
print("¡Los conjuntos \'x_test\' y \'y_test\' tienen las dimensiones correctas!")
```

Output:
```
¡Los conjuntos 'x_test' y 'y_test' tienen las dimensiones correctas!
```

### Ejercicio 4.1.

Realiza el encoding de las variables categóricas para el conjunto de test.


#### Ejercicio 4.1.1.

Define los nombres de las variables categóricas y numéricas.

<!-- grade_id: cell-df69761bb1ce4585, points: 0, grade: false, solution: true -->
```python
# Tu respuesta deben ser dos líneas consecutivas:
#    nombres_variables_categoricas = test.<<Función>>
#    nombres_variables_numericas = <<Consulta>>
# your code here
nombres_variables_categoricas = ["BusinessTravel", "Department", "EducationField", 
                                 "Gender", "JobRole", "MaritalStatus", "OverTime"]
nombres_variables_numericas = list(set(x_test.columns) - set(nombres_variables_categoricas))
```

<!-- grade_id: cell-f7366615c22fe2b1, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables existen
# La variable "nombres_variables_categoricas" tiene las columnas categóricas
# La variable "nombres_variables_numericas" tiene las columnas numéricas
# Las variables tienen las dimensiones correctas
#--------------------------------------
assert nombres_variables_categoricas is not None, "Asegúrate de definir la variable para columnas categóricas"
assert nombres_variables_numericas is not None, "Asegúrate de definir la variable para columnas numéricas"

assert len(nombres_variables_categoricas) > 0, "Asegúrate de escribir las columnas categóricas"
assert len(nombres_variables_numericas) > 0, "Asegúrate de escribir las columnas numéricas"

print("Separaste los tipos de las variables correctamente!")
```

Output:
```
Separaste los tipos de las variables correctamente!
```

#### Ejercicio 4.1.2.

Usa las dos variables anteriores para separar las columnas de los datos.

<!-- grade_id: cell-72ea9d8ef756c3ba, points: 0, grade: false, solution: true -->
```python
# Tu respuesta deben ser dos líneas consecutivas:
#    x_test_num = x_test[<<Consulta>>]
#    x_test_cat = x_test[<<Consulta>>]
# your code here
x_test_cat = x_test[nombres_variables_categoricas]
x_test_num = x_test[nombres_variables_numericas]
```

<!-- grade_id: cell-c8a79f521612450c, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables existen
# La variable "x_test_num" tiene los datos numéricos
# La variable "x_test_cat" tiene los datos categóricos
# Las variables tienen las dimensiones correctas
#--------------------------------------
assert x_test_cat is not None, "Asegúrate de definir la variable para datos categóricos"
assert x_test_num is not None, "Asegúrate de definir la variable para datos numéricos"

assert x_test_cat.shape == (294,7), "Asegúrate de separar los datos categóricos, no debería cambiar la cantidad de filas"
assert x_test_num.shape == (294,24), "Asegúrate de separar los datos numéricos, no debería cambiar la cantidad de filas"

print("Separaste los datos correctamente!")
```

Output:
```
Separaste los datos correctamente!
```

#### Ejercicio 4.1.3.

Transforma los datos categóricos usando el encoder que está en la variable `encoder`, no es necesario crear un encoder nuevo y ajustarlo.

<!-- grade_id: cell-073d33ae7f3d8e4b, points: 0, grade: false, solution: true -->
```python
# Tu respuesta debe tener la siguiente forma o cumplir el mismo propósito:
#    x_test_cat = encoder.<<Función>>
#    x_test_cat = pd.DataFrame(encoder.transform(x_copia_cat),columns=encoder.get_feature_names_out())
columnas = encoder.get_feature_names_out()
# your code here
x_test_cat = pd.DataFrame(encoder.transform(x_test_cat),columns=encoder.get_feature_names_out())
```

<!-- grade_id: cell-e0e3d6ddb19e1dc9, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variable existe
# La variable "x_test_num" tiene los datos numéricos
# La variable "x_test_cat" tiene los datos codificados
# Las variable tiene las dimensiones correctas
#--------------------------------------

assert x_test_cat is not None, "Asegúrate de definir la variable para datos categóricos"
assert columnas is not None, "No borres la variable con las columnas nuevas"

assert x_test_cat.shape == (294,28), "La cantidad de columnas debió aumentar, sin cambiar la cantidad de registros."

print("¡los datos se codificaron correctamente!")
```

Output:
```
¡los datos se codificaron correctamente!
```

#### Ejercicio 4.1.4.
Combina las los datos de `x_test_num` y `x_test_cat` en un único dataset en `x_test` usando pandas (***Ejemplo* `pd.<<Función>>`)

<!-- grade_id: cell-e8deb6a1b4662a9e, points: 0, grade: false, solution: true -->
```python
# your code here
x_test = pd.concat([x_test_num.reset_index(drop=True), x_test_cat.reset_index(drop=True)], axis=1)
x_test.head()
```

Output:
```
NumCompaniesWorked  TrainingTimesLastYear  DailyRate  JobLevel  \
0                   1                      3        635         3   
1                   1                      2        575         2   
2                   0                      0        662         2   
3                   4                      6       1492         3   
4                   1                      3        459         1   

   TotalWorkingYears  YearsInCurrentRole  StockOptionLevel  \
0                 10                   3                 0   
1                  5                   3                 0   
2                  5                   2                 1   
3                 14                  10                 1   
4                  1                   0                 0   

   YearsSinceLastPromotion  JobSatisfaction  PercentSalaryHike  ...  \
0                        9                4                 14  ...   
1                        0                2                 13  ...   
2 
```

<!-- grade_id: cell-68e08c2108d26ad9, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "x_test" es un DataFrame
# La variable tiene las dimensiones correctas
#--------------------------------------

# Se verifica que "x_test" sea un DataFrame
assert isinstance(x_test, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable
assert x_test.shape == (294,52), "\'x_test\' debe mantener las mismas dimensiones, solo se modifican sus valores."
print("¡\'x_test\' es un DataFrame con las dimensiones correctas!")
```

Output:
```
¡'x_test' es un DataFrame con las dimensiones correctas!
```

### Ejercicio 4.2.

Realiza la estandarización de los datos del conjunto de pruebas. Modifica el conjunto de variables independientes, `x_test`, con las operaciones necesarias para estandarizarlo. Particularmente, deberás seguir estos pasos:

* Utiliza la variable `scaler` para transformar los datos de `x_test` utilizando **solamente la información del conjunto de entrenamiento**. Asigna el resultado a la misma variable, sobreescribiendo su contenido. (**Ejemplo: `x_test = scaler.<<Función>>`**)
* Reconstruye el DataFrame de variables independientes utilizando la información almacenada en `columns`. Asigna el resultado a `x_test`. (**Ejemplo: `x_test = <<Función>>`**)

<!-- grade_id: E42, points: 0, grade: false, solution: true -->
```python
# your code here
columns = x_test.columns
scaler = StandardScaler()
x_test = scaler.fit_transform(x_test)
x_test = pd.DataFrame(x_test, columns=columns)
```

<!-- grade_id: E42-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "x_test" es un DataFrame
# La variable tiene las dimensiones correctas
#--------------------------------------

# Se verifica que "x_test" sea un DataFrame
assert isinstance(x_test, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable
assert x_test.shape == (294,52), "\'x_test\' debe mantener las mismas dimensiones, solo se modifican sus valores."
print("¡\'x_test\' es un DataFrame con las dimensiones correctas!")
```

Output:
```
¡'x_test' es un DataFrame con las dimensiones correctas!
```

### Ejercicio 4.3.

Con el conjunto de pruebas preparado, realiza predicciones con el fin de compararlas con los valores reales almacenados en `y_test`.

* Utiliza la variable `mejor_modelo` para realizar las predicciones sobre el mejor modelo. Asigna el resultado a una variable con nombre `y_pred` (**Ejemplo: `y_pred = mejor_modelo.<<Función>>`**).

<!-- grade_id: E43, points: 0, grade: false, solution: true -->
```python
# your code here
y_pred = mejor_modelo.predict(x_test)
```

<!-- grade_id: E43-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "y_pred" existe
# La variable "y_pred" es un arreglo
# La variable "y_pred" tiene las dimensiones correctas
#--------------------------------------

# Se verifica que "y_pred" está definida
assert y_pred is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "y_pred" sea un arreglo
assert isinstance(y_pred, np.ndarray), "El resultado debe ser un arreglo de Numpy."

# Se evalúan las dimensiones de "y_pred"
assert y_pred.shape == (294,), "\'y_pred\' debe tener el mismo número de filas (predicciones) que \'x_test\' y \'y_test\'."
print("¡\'y_pred\' es un arreglo con las dimensiones correctas!")
```

Output:
```
¡'y_pred' es un arreglo con las dimensiones correctas!
```

### Ejercicio 4.4.

Utiliza el conjunto de predicciones `y_pred` y el conjunto de valores reales `y_test` para obtener la matriz de confusión del mejor modelo.

* Haz un llamado a la función que retorna la matriz de confusión como un arreglo. Asigna el resultado a una nueva variable con el nombre `p44` (**Ejemplo: `p44 = <<Función>>`**).
* Encontrarás la línea `p44` al final de la celda. Esta línea se usa para que puedas visualizar la matriz resultante, por lo que no la debes modificar.

<!-- grade_id: E44, points: 0, grade: false, solution: true -->
```python
# your code here
p44 = confusion_matrix(y_test, y_pred) 

p44
```

Output:
```
array([[241,   4],
       [ 31,  18]])
```

<!-- grade_id: E44-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "p44" existe
# La variable "p44" es un arreglo de numpy
# La variable "p44" tiene las dimensiones correctas
#--------------------------------------

# Se verifica que "p44" está definida
assert p44 is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "p44" sea un arreglo de numpy
assert isinstance(p44, np.ndarray), "El resultado debe ser un arreglo de Numpy."

# Se evalúan las dimensiones de "p44"
assert p44.shape == (2,2), "\'p44\' debe tener dos filas y dos columnas, equivalente al número de clases del problema."
print("¡\'p44\' es un arreglo con las dimensiones correctas!")
```

Output:
```
¡'p44' es un arreglo con las dimensiones correctas!
```

### Ejercicio 4.5.

Finalmente, obtén cuatro métricas de rendimiento utilizando las predicciones del modelo.

#### Ejercicio 4.5.1.

Utiliza `scikit-learn` para obtener la exactitud (accuracy).

* Define una variable con el nombre `accuracy` y asígnale la función necesaria para obtener la exactitud (`accuracy = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E451, points: 0, grade: false, solution: true -->
```python
# your code here
accuracy = accuracy_score(y_test, y_pred)
print("Exactitud: ", accuracy)
```

Output:
```
Exactitud:  0.8809523809523809
```

<!-- grade_id: E451-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "accuracy" existe
# La variable "accuracy" es un número
# La variable "accuracy" tiene un valor válido
#--------------------------------------

# Se verifica que "accuracy" está definida
assert accuracy is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "accuracy" sea un número
assert isinstance(accuracy, np.float64) or isinstance(accuracy, float), "El resultado debe ser un número decimal."

# Se verifica que "accuracy" tiene un valor válido
assert accuracy >= 0 and accuracy <= 1, "\'accuracy\' debe tener un valor entre 0 y 1."
print("¡\'accuracy\' es un número válido!")
```

Output:
```
¡'accuracy' es un número válido!
```

#### Ejercicio 4.5.2.

Utiliza `scikit-learn` para obtener la sensibilidad (recall).

* Define una variable con el nombre `recall` y asígnale la función necesaria para obtener la sensibilidad (`recall = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E452, points: 0, grade: false, solution: true -->
```python
# your code here
recall = recall_score(y_test,y_pred)
print("Sensibilidad: ", recall)
```

Output:
```
Sensibilidad:  0.3673469387755102
```

<!-- grade_id: E452-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "recall" existe
# La variable "recall" es un número
# La variable "recall" tiene un valor válido
#--------------------------------------

# Se verifica que "recall" está definida
assert recall is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "recall" sea un número
assert isinstance(recall, np.float64) or isinstance(recall, float), "El resultado debe ser un número decimal."

# Se verifica que "recall" tiene un valor válido
assert recall >= 0 and recall <= 1, "\'recall\' debe tener un valor entre 0 y 1."
print("¡\'recall\' es un número válido!")
```

Output:
```
¡'recall' es un número válido!
```

#### Ejercicio 4.5.3.

Utiliza `scikit-learn` para obtener la precisión.

* Define una variable con el nombre `precision` y asígnale la función necesaria para obtener la precisión (`precision = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E453, points: 0, grade: false, solution: true -->
```python
# your code here
precision = precision_score(y_test,y_pred)
print('Precisión: ', precision)
```

Output:
```
Precisión:  0.8181818181818182
```

<!-- grade_id: E453-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "precision" existe
# La variable "precision" es un número
# La variable "precision" tiene un valor válido
#--------------------------------------

# Se verifica que "precision" está definida
assert precision is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "precision" sea un número
assert isinstance(precision, np.float64) or isinstance(precision, float), "El resultado debe ser un número decimal."

# Se verifica que "precision" tiene un valor válido
assert precision >= 0 and precision <= 1, "\'precision\' debe tener un valor entre 0 y 1."
print("¡\'precision\' es un número válido!")
```

Output:
```
¡'precision' es un número válido!
```

#### Ejercicio 4.5.4.

Utiliza `scikit-learn` para obtener el F1-Score.

* Define una variable con el nombre `f1` y asígnale la función necesaria para obtener el F1-Score (`f1 = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E454, points: 0, grade: false, solution: true -->
```python
# your code here
f1 = f1_score(y_test,y_pred)
print('F1-Score: ', f1)
```

Output:
```
F1-Score:  0.5070422535211268
```

<!-- grade_id: E454-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "f1" existe
# La variable "f1" es un número
# La variable "f1" tiene un valor válido
#--------------------------------------

# Se verifica que "f1_score" está definida
assert f1 is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "f1_score" sea un número
assert isinstance(f1, np.float64) or isinstance(f1, float), "El resultado debe ser un número decimal."

# Se verifica que "f1_score" tiene un valor válido
assert f1 >= 0 and f1 <= 1, "\'f1_score\' debe tener un valor entre 0 y 1."
print("¡\'f1_score\' es un número válido!")
```

Output:
```
¡'f1_score' es un número válido!
```

## Cierre

Al realizar los ejercicios de este taller, has reforzado tus capacidades para codificar variables categóricas usando la codificación One Hot, realizar una búsqueda exhaustiva de hiperparámetros y evaluar el mejor modelo de regresión logística resultante utilizando la matriz de confusión y cuatro métricas de rendimiento. 

---
*Creado por: Nicolás Díaz*

*Última edición: Camilo Rozo*

*Revisado por: Haydemar Nuñez*

*Versión: Febrero 2025*  

*Universidad de los Andes*   