# Un problema de regresión lineal

En este taller estudiarás conceptos mostrados en el tutorial de "Solución de un problema de regresión lineal", utilizando el conjunto de datos de popularidad de canciones. Particularmente, realizarás los siguientes procesos:

1. Cargar un conjunto de datos en formato .csv.
3. Preparar los datos para el modelado.
4. Entrenar un modelo de regresión lineal.
5. Evaluar el modelo resultante.

El problema a resolver es el siguiente: dadas las propiedades acústicas de una canción, predecir su popularidad. Antes de iniciar, es necesario importar algunas librerías: 

```python
import pandas as pd
import numpy as np

from sklearn.model_selection import train_test_split
from sklearn.linear_model import LinearRegression
from sklearn.metrics import root_mean_squared_error,mean_squared_error, mean_absolute_error, r2_score

from importlib.metadata import version
```

También es útil revisar las versiones de las librerías

```python
print(f"Versión de Pandas: {version('pandas')}")
print(f"Versión de Scikit-learn: {version('scikit-learn')}")
print(f"Versión de Numpy: {version('numpy')}")
```

Output:
```
Versión de Pandas: 2.2.2
Versión de Scikit-learn: 1.7.1
Versión de Numpy: 1.26.4
```

## 1. Carga de datos

Con las librerías importadas, seguiremos con la carga del conjunto de datos.

### Ejercicio 1.1.

Utiliza Pandas para importar el archivo que contiene el conjunto de datos de canciones.

* La ruta del archivo .csv es: `./data/song_data.csv`, y ya se encuentra en el entorno de Coursera, solo debes importarlo.
* La variable resultante debe tener el nombre `data_raw`, que representa el conjunto de datos sin modificar.

<!-- grade_id: E11, points: 0, grade: false, solution: true -->
```python
ruta = './data/song_data.csv'
# your code here
data_raw = pd.read_csv(ruta)
```

<!-- grade_id: E11-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "data_raw" existe
# El resultado es un DataFrame
# El resultado tiene las dimensiones correctas
#--------------------------------------

# Se verifica que la variable está definida
assert data_raw is not None, "Asegúrate de definir la variable correctamente."
assert ruta is not None, "Asegúrate de definir la ruta de los datos correctamente."

# Se verifica que sea un DataFrame
assert isinstance(data_raw, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable data
assert data_raw.shape == (18835,15), "¿Verificaste que la ruta del archivo CSV y el nombre de la variable son correctos?"
print("¡Los datos tienen las dimensiones correctas!")
```

Output:
```
¡Los datos tienen las dimensiones correctas!
```

En este punto puedes agregar una celda y verificar el tamaño de los datos y su estructura, como último paso antes de la preparación del conjunto.

```python
# Verificar las dimensiones del DataFrame:
data_raw.shape
```

Output:
```
(18835, 15)
```

```python
# Revisar las primeras filas del DataFrame para asegurarse de que se cargó correctamente:
data_raw.head()
```

Output:
```
song_name  song_popularity  song_duration_ms  \
0  Boulevard of Broken Dreams               73            262333   
1                  In The End               66            216933   
2           Seven Nation Army               76            231733   
3                  By The Way               74            216933   
4           How You Remind Me               56            223826   

   acousticness  danceability  energy  instrumentalness  key  liveness  \
0      0.005520         0.496   0.682          0.000029    8    0.0589   
1      0.010300         0.542   0.853          0.000000    3    0.1080   
2      0.008170         0.737   0.463          0.447000    0    0.2550   
3      0.026400         0.451   0.970          0.003550    0    0.1020   
4      0.000954         0.447   0.766          0.000000   10    0.1130   

   loudness  audio_mode  speechiness    tempo  time_signature  audio_valence  
0    -4.095           1       0.0294  167.060               4          0.474  
1    -6.
```

## 2. Preparación de los datos

Utilizaremos una nueva variable `data` para verificar que el conjunto no contiene errores, además de realizar las modificaciones necesarias al conjunto de datos. Recuerda que usarás esta variable para definir los conjuntos de entrenamiento y de pruebas.

```python
data = data_raw.copy()
```

### Ejercicio 2.1.

Utiliza Pandas para obtener la cantidad de datos faltantes por variable.

* Define una variable con el nombre `p21` y asígnale las operaciones correspondientes para obtener la lista de datos faltantes. (`p21 = <<Operaciones>>`)
* La respuesta debe ser un arreglo con los datos faltantes por variable. Este arreglo es resultado de aplicar dos funciones a la variable `data`. (**Ejemplo: `p21 = data.<<Función1>>.<<Función2>>`**)
* Encontrarás una línea solo con el nombre de la variable. Esta línea se usa para que puedas visualizar tu respuesta, por lo que siempre debe ir al final y no la debes modificar.

<!-- grade_id: E21, points: 0, grade: false, solution: true -->
```python
# your code here
# Lista de datos faltantes por columna:
p21 = data.isna().sum()
p21
```

Output:
```
song_name           0
song_popularity     0
song_duration_ms    0
acousticness        0
danceability        0
energy              0
instrumentalness    0
key                 0
liveness            0
loudness            0
audio_mode          0
speechiness         0
tempo               0
time_signature      0
audio_valence       0
dtype: int64
```

<!-- grade_id: E21-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "p21" existe
# El resultado es una Serie de Pandas
# El resultado tiene las dimensiones correctas
#--------------------------------------

# Se verifica que la variable está definida
assert p21 is not None, "Asegúrate de definir la variable correctamente."

# Se verifica que sea una Serie de Pandas
assert isinstance(p21, pd.Series), "El resultado debe ser una Serie de Pandas."

# Se verifican las dimensiones de la respuesta
assert p21.shape == (15,), "Utiliza una función de Pandas para obtener los datos faltantes, y luego usa otra función para sumarlos."
print("¡Las dimensiones son correctas!")
```

Output:
```
¡Las dimensiones son correctas!
```

### Ejercicio 2.2.

Similar al ejercicio anterior, utiliza Pandas para obtener la cantidad de filas duplicadas.

* Define una variable con el nombre `p22` y asígnale las operaciones correspondientes para obtener el número de filas duplicadas (`p22 = <<Operaciones>>`).
* La respuesta debe ser un número entero con la cantidad de datos (filas) duplicados. Este número es resultado de aplicar dos funciones consecutivas a la variable `data`. (**Ejemplo: `p22 = data.<<Función1>>.<<Función2>>`**)
* Encontrarás una línea solo con el nombre de la variable. Esta línea se usa para que puedas visualizar tu respuesta, por lo que siempre debe ir al final y no la debes modificar.

<!-- grade_id: E22, points: 0, grade: false, solution: true -->
```python
# your code here
# Número de filas duplicadas en el DataFrame:
p22 = data.duplicated().sum()

p22
```

Output:
```
3909
```

<!-- grade_id: E22-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "p22" existe
# El resultado es un número entero
# El resultado está en el rango [0,18835]
#--------------------------------------

# Se verifica que la variable está definida
assert p22 is not None, "Asegúrate de definir la variable correctamente."

# Se evalúa el tipo de dato
assert type(p22) == np.int64 or type(p22) == int, "Recuerda que tu respuesta debe ser un número entero."

# Se evalúa el rango de respuesta
assert (p22 >= 0 and p22 <= 18835), "El número de filas duplicadas debe ser positivo y menor al total de filas"
print("¡El tipo y rango de la respuesta son correctos!")
```

Output:
```
¡El tipo y rango de la respuesta son correctos!
```

Es conveniente eliminar los datos duplicados, por lo que utilizaremos la función `drop_duplicates()`:

```python
data = data.drop_duplicates()
```

Después de la limpieza de datos, puedes crear una celda para verificar el tamaño del conjunto resultante.

```python
data.shape
```

Output:
```
(14926, 15)
```

### Ejercicio 2.3.

Teniendo en cuenta que queremos predecir la popularidad de una canción dadas sus características acústicas, vamos a eliminar una variable que no es muy relevante: el nombre. A continuación, utiliza Pandas para eliminar la columna `song_name` de la variable `data`.

* Para este punto no debes definir una nueva variable, sino asignar tu respuesta a la variable `data`. (**Ejemplo: `data = data.<<Función>>`**)

<!-- grade_id: E23, points: 0, grade: false, solution: true -->
```python
# your code here
# Eliminamos la variable song_name poco relevante para el modelo de regresión lineal:
data = data.drop(columns=['song_name'], axis="columns")
```

<!-- grade_id: E23-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "data" existe
# El resultado es un DataFrame
# El resultado tiene las dimensiones correctas
# El resultado no contiene la variable "song_name"
#--------------------------------------

# Se verifica que la variable está definida
assert data is not None, "Asegúrate de definir la variable correctamente."

# Se verifica que sea un DataFrame
assert isinstance(data, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se evalúan las dimensiones de la variable data
assert data.shape == (14926,14), "¿Verificaste que el nombre de la variable es correcto y que la función utilizada es correcta?"

# Se verifica que la variable eliminada es correcta
assert "song_name" not in data.columns, "Recuerda que debes eliminar la columna \'song_name\'."
print("¡La variable tiene las dimensiones correctas!")
```

Output:
```
¡La variable tiene las dimensiones correctas!
```

## 3. Entrenamiento del modelo

Con los datos preparados, ahora entrenarás el modelo de regresión lineal.

### Ejercicio 3.1.

Inicialmente, tenemos que dividir el conjunto de datos en entrenamiento y pruebas. Usando el 80% de los datos para entrenar el modelo y el 20% restante para probarlo, utiliza `scikit-learn` para separar el conjunto de datos en dos.

* Guarda tu respuesta en dos variables: `train` y `test`. (**Ejemplo: `train, test = <<Función>>`**)
* Utiliza el parámetro `random_state=0`. Esto hará que la partición sea siempre la misma.
* Encontrarás la línea `train.head()`. Esta línea se usa para que puedas visualizar el resultado del conjunto de entrenamiento. Déjala al final de la celda y no la modifiques.

<!-- grade_id: E31, points: 0, grade: false, solution: true -->
```python
# your code here
# Dividir el DataFrame en conjuntos de entrenamiento y prueba (80% y 20% respectivamente):
train, test = train_test_split(data, test_size=0.2, random_state=0)

train.head()
```

Output:
```
song_popularity  song_duration_ms  acousticness  danceability  energy  \
1154                68            214880       0.00818         0.527   0.941   
13027               59            237080       0.07750         0.658   0.862   
13150               33            213211       0.06630         0.813   0.777   
11215               46            269096       0.97200         0.492   0.240   
13523               56            189500       0.45800         0.773   0.725   

       instrumentalness  key  liveness  loudness  audio_mode  speechiness  \
1154           0.000002    6    0.3970    -3.376           1       0.0551   
13027          0.000000    6    0.1230    -5.992           0       0.0445   
13150          0.042100    5    0.0785    -4.355           0       0.0828   
11215          0.141000   11    0.0736   -13.386           1       0.0397   
13523          0.000000    8    0.0935    -5.614           1       0.0344   

         tempo  time_signature  audio_valence  
1154    94.056 
```

<!-- grade_id: E31-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables "train" y "test" existen
# Las variables "train" y "test" son un DataFrame
# Las variables tienen las dimensiones correctas
#--------------------------------------

# Se verifica que las variables están definidas
assert train is not None, "Asegúrate de definir la variable correctamente."
assert test is not None, "Asegúrate de definir la variable correctamente."

# Se verifica que "train" y "test" sean un DataFrame
assert isinstance(train, pd.DataFrame), "La variable \'train\' debe ser un DataFrame."
assert isinstance(test, pd.DataFrame), "La variable \'test\' debe ser un DataFrame."

# Se evalúan las dimensiones de las variables
assert train.shape == (11940,14), "Verifica que estés utilizando el 80% de los datos para el conjunto de entrenamiento."
assert test.shape == (2986,14), "Verifica que estés utilizando el 20% de los datos para el conjunto de pruebas."
print("¡Los conjuntos de entrenamiento y pruebas tienen las dimensiones correctas!")
```

Output:
```
¡Los conjuntos de entrenamiento y pruebas tienen las dimensiones correctas!
```

### Ejercicio 3.2.

Ahora trabajaremos sobre el conjunto de entrenamiento. Primero debes aislar la variable objetivo, `song_popularity`, de las variables independientes. Utiliza Pandas para crear dos variables, `x_train` y `y_train`, que almacenarán las variables independientes y la variable objetivo, respectivamente.

* Crea una variable con nombre `x_train` y asígnale la operación necesaria para almacenar solo las variables independientes del conjunto de entrenamiento. (**Ejemplo: `x_train = train.<<Función>>`**)
* Crea una variable con nombre `y_train` y asígnale la operación necesaria para almacenar la variable objetivo del conjunto de entrenamiento. (**Ejemplo: `y_train = <<Consulta>>`**)

<!-- grade_id: E32, points: 0, grade: false, solution: true -->
```python
# your code here
# Aislamos las variables predictoras y la variable objetivo para el conjunto de entrenamiento:
x_train = train.drop(columns=['song_popularity'], axis=1)
y_train = train['song_popularity']

```

<!-- grade_id: E32-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables "x_train" y "y_train" existen
# La variable "x_train" es un DataFrame
# La variable "y_train" es una Serie de Pandas
# Las variables tienen las dimensiones correctas
#--------------------------------------

# Se verifica que las variables están definidas
assert x_train is not None, "Asegúrate de definir la variable correctamente."
assert y_train is not None, "Asegúrate de definir la variable correctamente."

# Se verifica que "x_train" sea un DataFrame
assert isinstance(x_train, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se verifica que "y_train" sea una Serie
assert isinstance(y_train, pd.Series), "El resultado debe ser una Serie de Pandas."

# Se evalúan las dimensiones de las variables
assert x_train.shape == (11940,13), "\'x_train\' debe tener el mismo número de filas pero una columna menos que \'train\'."
assert y_train.shape == (11940,), "\'y_train\' solamente contiene una columna."
print("¡Los conjuntos \'x_train\' y \'y_train\' tienen las dimensiones correctas!")
```

Output:
```
¡Los conjuntos 'x_train' y 'y_train' tienen las dimensiones correctas!
```

En este punto, es útil verificar que `x_train` no contiene la variable objetivo, y que `y_train` solamente contiene la popularidad de las canciones.

```python
# Variables predictoras para el conjunto de entrenamiento:
x_train.head()
```

Output:
```
song_duration_ms  acousticness  danceability  energy  instrumentalness  \
1154             214880       0.00818         0.527   0.941          0.000002   
13027            237080       0.07750         0.658   0.862          0.000000   
13150            213211       0.06630         0.813   0.777          0.042100   
11215            269096       0.97200         0.492   0.240          0.141000   
13523            189500       0.45800         0.773   0.725          0.000000   

       key  liveness  loudness  audio_mode  speechiness    tempo  \
1154     6    0.3970    -3.376           1       0.0551   94.056   
13027    6    0.1230    -5.992           0       0.0445  149.980   
13150    5    0.0785    -4.355           0       0.0828  125.065   
11215   11    0.0736   -13.386           1       0.0397   85.490   
13523    8    0.0935    -5.614           1       0.0344   99.018   

       time_signature  audio_valence  
1154                4          0.440  
13027               4          0.
```

```python
# Variable objetivo para el conjunto de entrenamiento:
y_train.head()
```

Output:
```
1154     68
13027    59
13150    33
11215    46
13523    56
Name: song_popularity, dtype: int64
```

### Ejercicio 3.3.

Ahora podemos entrenar el modelo con los datos modificados. Inicialmente crearemos un objeto de tipo `LinearRegression()`:

```python
regresion = LinearRegression()
```

A continuación, entrena el modelo utilizando el conjunto de entrenamiento, compuesto por las variables `x_train` y `y_train`.

* Para este ejercicio no debes asignar tu resultado a ninguna variable. Es decir, solo debes ejecutar una función utilizando las variables `x_train` y `y_train` como parámetros. (**Ejemplo: `regresion.<<Función>>`**)

<!-- grade_id: E33, points: 0, grade: false, solution: true -->
```python
# your code here
# Entrenamos el modelo de regresión lineal con los datos de entrenamiento:
regresion.fit(x_train, y_train)
```

Output:
```
LinearRegression()
```

<!-- grade_id: E33-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# El atributo "coef_" y el atributo "intercept_" existen
# El objeto "regresion" tiene 13 coeficientes
# El objeto "regresion" tiene un intercepto diferente de cero
#--------------------------------------

# Se verifica que los atributos están definidos
assert regresion.coef_ is not None, "Asegúrate de definir la variable correctamente."
assert regresion.intercept_ is not None, "Asegúrate de definir la variable correctamente."

# Se verifica que "regresion" haya generado coeficientes
assert regresion.coef_.shape == (13,), "Al ejecutar una función usando el objeto \'regresion\', se deben generar 13 coeficientes. Verifica que ejecutaste la función para entrenar el modelo."

# Se verifica que "regresion" haya generado un intercepto
assert regresion.intercept_ > 0, "Al ejecutar una función usando el objeto \'regresion\', se debe generar un intercepto. Verifica que ejecutaste la función para entrenar el modelo."
print("¡Se ha entrenado el modelo correctamente!")
```

Output:
```
¡Se ha entrenado el modelo correctamente!
```

El resultado del entrenamiento son los coeficientes y el intercepto que definen una recta. Puedes observarlos a continuación:

```python
print ('Coeficientes:')
display(pd.DataFrame(zip(x_train.columns, regresion.coef_),columns=["Variable","Coeficiente"]))
print ('Intercepto: ',regresion.intercept_)
```

Output:
```
Coeficientes:
            Variable  Coeficiente
0   song_duration_ms    -0.000002
1       acousticness    -2.630871
2       danceability     6.067543
3             energy   -10.377233
4   instrumentalness    -6.611366
5                key     0.011543
6           liveness    -3.955941
7           loudness     0.500937
8         audio_mode     0.555595
9        speechiness    -3.428583
10             tempo    -0.009264
11    time_signature     0.712617
12     audio_valence    -5.452651
Intercepto:  58.93734587158678
```

## 4. Evaluación del modelo

Por último, evaluarás el modelo entrenado utilizando el conjunto de pruebas.

### Ejercicio 4.1.

Inicialmente, separa las variables independientes y la variable objetivo en el conjunto de pruebas. Para ello, utiliza Pandas para crear dos variables, `x_test` y `y_test`, que almacenarán las variables independientes y la variable objetivo, respectivamente.

* Crea una variable con nombre `x_test` y asígnale la operación necesaria para almacenar solo las variables independientes del conjunto de pruebas. (**Ejemplo: `x_test = test.<<Función>>`**)
* Crea una variable con nombre `y_test` y asígnale la operación necesaria para almacenar la variable objetivo del conjunto de pruebas. (**Ejemplo: `y_test = <<Consulta>>`**)

<!-- grade_id: E41, points: 0, grade: false, solution: true -->
```python
# Tu respuesta deben ser dos líneas consecutivas:
#    x_test = test.<<Función>>
#    y_test = <<Consulta>>
# your code here
# Dividimos el conjunto de prueba en variables predictoras y variable objetivo:
x_test = test.drop(columns=['song_popularity'], axis="columns")
y_test = test['song_popularity']
```

<!-- grade_id: E41-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# Las variables "x_test" y "y_test" existen
# La variable "x_test" es un DataFrame
# La variable "y_test" es una Serie de Pandas
# Las variables tienen las dimensiones correctas
#--------------------------------------

# Se verifica que las variables están definidas
assert x_test is not None, "Asegúrate de definir la variable correctamente."
assert y_test is not None, "Asegúrate de definir la variable correctamente."

# Se verifica que "x_test" sea un DataFrame
assert isinstance(x_test, pd.DataFrame), "El resultado debe ser un DataFrame."

# Se verifica que "y_test" sea una Serie
assert isinstance(y_test, pd.Series), "El resultado debe ser una Serie de Pandas."

# Se evalúan las dimensiones de las variables
assert x_test.shape == (2986,13), "\'x_test\' debe tener el mismo número de filas pero una columna menos que \'test\'."
assert y_test.shape == (2986,), "\'y_test\' solamente contiene una columna."
print("¡Los conjuntos \'x_test\' y \'y_test\' tienen las dimensiones correctas!")
```

Output:
```
¡Los conjuntos 'x_test' y 'y_test' tienen las dimensiones correctas!
```

### Ejercicio 4.2.

Con el conjunto de pruebas preparado, realiza predicciones con el fin de compararlas con los valores reales almacenados en `y_test`.

* Utiliza el objeto `regresion` para realizar las predicciones. Asigna el resultado a una variable con nombre `y_pred` (**Ejemplo: `y_pred = regresion.<<Función>>`**).
* Encontrarás una línea solo con el nombre de la variable. Esta línea se usa para que puedas visualizar tu respuesta, por lo que siempre debe ir al final y no la debes modificar.

<!-- grade_id: E42, points: 0, grade: false, solution: true -->
```python
# your code here
# Realizamos las predicciones del modelo de regresión lineal con los datos de prueba:
y_pred = regresion.predict(x_test)
y_pred
```

Output:
```
array([48.628855  , 43.25413623, 49.27517053, ..., 51.85949766,
       50.20594692, 46.88693145])
```

<!-- grade_id: E42-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "y_pred" existe
# La variable "y_pred" es un arreglo
# La variable "y_pred" tiene las dimensiones correctas
#--------------------------------------

# Se verifica que la variable está definida
assert y_pred is not None, "Asegúrate de definir la variable correctamente."

# Se verifica que "y_pred" sea un arreglo
assert isinstance(y_pred, np.ndarray), "El resultado debe ser un Arreglo de Numpy."

# Se evalúan las dimensiones de "y_pred"
assert y_pred.shape == (2986,), "\'y_pred\' debe tener el mismo número de filas (predicciones) que \'x_test\' y \'y_test\'."
print("¡\'y_pred\' es un arreglo con las dimensiones correctas!")
```

Output:
```
¡'y_pred' es un arreglo con las dimensiones correctas!
```

### Ejercicio 4.3.

Finalmente, obtén tres métricas de rendimiento utilizando las predicciones del modelo.

#### Ejercicio 4.3.1.

Utiliza `scikit-learn` para obtener la Raíz del Error Cuadrático Medio (RMSE).

* Define una variable con el nombre `rmse` y asígnale la función necesaria para obtener el RMSE (`rmse = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E431, points: 0, grade: false, solution: true -->
```python
# your code here
# Calculamos el RMSE (Root Mean Squared Error) para evaluar el desempeño del modelo de regresión lineal:
rmse = root_mean_squared_error(y_test, y_pred)
print("RMSE: ", rmse)
```

Output:
```
RMSE:  20.28462756391579
```

<!-- grade_id: E431-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "rmse" existe
# La variable "rmse" es un número
# La variable "rmse" tiene un valor válido
#--------------------------------------

# Se verifica que "rmse" está definida
assert rmse is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "rmse" sea un número
assert isinstance(rmse, np.float64) or isinstance(rmse, float), "El resultado debe ser un número decimal."

# Se verifica que "rmse" tiene un valor válido
assert rmse >= 0 and rmse <= 100, "\'rmse\' debe tener un valor entre 0 y 100 (los valores posibles de la popularidad)."
print("¡\'rmse\' es un número válido!")
```

Output:
```
¡'rmse' es un número válido!
```

#### Ejercicio 4.3.2.

Utiliza `scikit-learn` para obtener el Error Absoluto Medio (MAE).

* Define una variable con el nombre `mae` y asígnale la función necesaria para obtener el MAE (`mae = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E432, points: 0, grade: false, solution: true -->
```python
# your code here
# Calculamos el MAE (Mean Absolute Error) para evaluar el desempeño del modelo de regresión lineal:
mae = mean_absolute_error(y_test, y_pred)
print("MAE: ", mae)
```

Output:
```
MAE:  16.281255362592358
```

<!-- grade_id: E432-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "mae" existe
# La variable "mae" es un número
# La variable "mae" tiene un valor válido
#--------------------------------------

# Se verifica que "mae" está definida
assert mae is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "mae" sea un número
assert isinstance(mae, np.float64) or isinstance(mae, float), "El resultado debe ser un número decimal."

# Se verifica que "mae" tiene un valor válido
assert mae >= 0 and mae <= 100, "\'mae\' debe tener un valor entre 0 y 100 (los valores posibles de la popularidad)."
print("¡\'mae\' es un número válido!")
```

Output:
```
¡'mae' es un número válido!
```

#### Ejercicio 4.3.3.

Utiliza `scikit-learn` para obtener el Coeficiente de determinación (R<sup>2</sup>).

* Define una variable con el nombre `r2` y asígnale la función necesaria para obtener el R<sup>2</sup> (`r2 = <<Función>>`).
* Encontrarás una línea que mostrará tu resultado. Esta línea debe ir al final y no la debes modificar.

<!-- grade_id: E433, points: 0, grade: false, solution: true -->
```python
# your code here
# Calculamos el R² (R-squared) para evaluar el desempeño del modelo de regresión lineal:
r2 = r2_score(y_test, y_pred)
print('R²: ', r2)
```

Output:
```
R²:  0.025595824929419142
```

<!-- grade_id: E433-tests, points: 1, grade: true, solution: false -->
```python
#---------- Celda de Pruebas ----------
# La variable "r2" existe
# La variable "r2" es un número
# La variable "r2" tiene un valor válido
#--------------------------------------

# Se verifica que "r2" está definida
assert r2 is not None, "Asegúrate de definir la variable con el nombre correcto."

# Se verifica que "r2" sea un número
assert isinstance(r2, np.float64) or isinstance(r2, float), "El resultado debe ser un número decimal."

# Se verifica que "r2" tiene un valor válido
assert r2 >= 0 and r2 <= 1, "\'r2\' debe tener un valor entre 0 y 1."
print("¡\'r2\' es un número válido!")
```

Output:
```
¡'r2' es un número válido!
```

## Cierre

Al desarrollar los ejercicios de este taller, has practicado tus capacidades para entrenar un modelo de regresión lineal mediante scikit-learn. Adicionalmente, realizaste predicciones con el modelo resultante y calculaste tres métricas de rendimiento sobre los resultados.

---
*Creado por: Nicolás Díaz*

*Última edición: Camilo Rozo*

*Revisado por: Haydemar Nuñez*

*Desarrollado por: María J. García-Bonilla*

*Versión de: Enero 20, 2025*  

*Universidad de los Andes*  